load("bf4b12814bc95f34eeb130127d8438ab.js");
load("93fae755edd261212639eed30afa2ca4.js");
// Copyright (C) 2016 Tim Disney. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: sec-template-literal-lexical-components
description: Invalid unicode escape sequence
negative:
  phase: early
  type: SyntaxError
---*/

`\u{10FFFFF}${'inner'}right`;
